# Basic linear algebra in Julia
Author: Andreas Noack Jensen (MIT & JuliaComputing) (https://twitter.com/anoackjensen?lang=en)
(with edits from Jane Herriman)

First let's define a random matrix

```julia
A = rand(1:4, 3, 3)
```

Output:
```
3×3 Matrix{Int64}:
 1  4  4
 3  3  3
 2  1  3
```

```julia
B = A
C = copy(A)
[B C]
```

Output:
```
3×6 Matrix{Int64}:
 1  4  4  1  4  4
 3  3  3  3  3  3
 2  1  3  2  1  3
```

```julia
A[1] = 17
[B C]
```

Output:
```
3×6 Matrix{Int64}:
 17  4  4  1  4  4
  3  3  3  3  3  3
  2  1  3  2  1  3
```

Define a vector of ones

```julia
x = fill(1.0, (3,)) # = fill(1.0, 3)
```

Output:
```
3-element Vector{Float64}:
 1.0
 1.0
 1.0
```

Notice that $A$ has type Array{Int64,2} but $x$ has type Array{Float64,1}. Julia defines the aliases Vector{Type}=Array{Type,1} and Matrix{Type}=Array{Type,2}. 

Many of the basic operations are the same as in other languages
#### Multiplication

```julia
b = A*x
```

Output:
```
3-element Vector{Float64}:
 25.0
  9.0
  6.0
```

#### Transposition
As in other languages `A'` is the conjugate transpose, or adjoint

```julia
A'
```

Output:
```
3×3 adjoint(::Matrix{Int64}) with eltype Int64:
 17  3  2
  4  3  1
  4  3  3
```

and we can get the transpose with

```julia
transpose(A)
```

Output:
```
3×3 transpose(::Matrix{Int64}) with eltype Int64:
 17  3  2
  4  3  1
  4  3  3
```

#### Transposed multiplication
Julia allows us to write this without *

```julia
A'A
```

Output:
```
3×3 Matrix{Int64}:
 302  79  83
  79  26  28
  83  28  34
```

#### Solving linear systems 
The problem $Ax=b$ for ***square*** $A$ is solved by the \ function.

```julia
A\b
```

Output:
```
3-element Vector{Float64}:
 1.0
 1.0000000000000004
 1.0
```

`A\b` gives us the *least squares solution* if we have an overdetermined linear system (a "tall" matrix)

```julia
Atall = rand(3, 2)
```

Output:
```
3×2 Matrix{Float64}:
 0.110558  0.395074
 0.32861   0.772552
 0.436232  0.529045
```

```julia
Atall\b
```

Output:
```
2-element Vector{Float64}:
 -47.07722501768563
  43.45101201297879
```

and the *minimum norm least squares solution* if we have a rank-deficient least squares problem

```julia
v = rand(3)
rankdef = hcat(v, v)
```

Output:
```
3×2 Matrix{Float64}:
 0.607886  0.607886
 0.850117  0.850117
 0.586299  0.586299
```

```julia
rankdef\b
```

Output:
```
2-element Vector{Float64}:
 9.180551036079004
 9.180551036079
```

Julia also gives us the minimum norm solution when we have an underdetermined solution (a "short" matrix)

```julia
bshort = rand(2)
Ashort = rand(2, 3)
```

Output:
```
2×3 Matrix{Float64}:
 0.692218  0.193236   0.153774
 0.127852  0.0359394  0.063034
```

```julia
Ashort\bshort
```

Output:
```
3-element Vector{Float64}:
 -4.139674407706176
 -0.9797338415813472
 23.563074056363842
```

# The LinearAlgebra library

While much of linear algebra is available in Julia by default (as shown above), there's a standard library named `LinearAlgebra` that brings in many more relevant names and functions. In particular, it provides factorizations and some structured matrix types.  As with all packages, you can bring these additional features into your session with a `using LinearAlgebra`.

### Exercises

#### 10.1 
Take the inner product (or "dot" product) of a vector `v` with itself and assign it to variable `dot_v`.



```julia
v = [1,2,3]
```

Output:
```
3-element Vector{Int64}:
 1
 2
 3
```

```julia
using Pkg
Pkg.add("LinearAlgebra")
using LinearAlgebra
```

Output:
```
[32m[1m   Resolving[22m[39m package versions...
[32m[1m    Updating[22m[39m `~/Documents/GitHub/Introduction-to-Julia/Project.toml`
  [90m[37e2e46d] [39m[92m+ LinearAlgebra[39m
[32m[1m  No Changes[22m[39m to `~/Documents/GitHub/Introduction-to-Julia/Manifest.toml`
```

```julia
dot_v = dot(v, v)
```

Output:
```
14
```

<!-- grade_id: cell-913fef9b0d19cd52, points: 1, grade: true, solution: false -->
```julia
@assert dot_v == 14
```

#### 10.2 
Take the outer product of a vector v with itself and assign it to variable `outer_v`

```julia
outer_v = v * v'
```

Output:
```
3×3 Matrix{Int64}:
 1  2  3
 2  4  6
 3  6  9
```

```julia
@assert outer_v == [1 2 3
                    2 4 6
                    3 6 9]
```

#### 10.3 
Use [LinearAlgebra.cross](https://docs.julialang.org/en/v1/stdlib/LinearAlgebra/#LinearAlgebra.cross) to compute the cross product of a vector v with itself and assign it to variable `cross_v`

```julia
cross_v = cross(v, v)
```

Output:
```
3-element Vector{Int64}:
 0
 0
 0
```

<!-- grade_id: cell-e6b6970ffe104df5, points: 1, grade: true, solution: false -->
```julia
@assert cross_v == [0, 0, 0]
```